<a href="https://colab.research.google.com/github/tamnhu2913/Applied-Social-Network-Analysis/blob/main/assignment3.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment 3

In this assignment you will explore measures of centrality on two networks, a friendship network in Part 1, and a blog network in Part 2.

## Part 1

Answer questions 1-4 using the network `G1`, a network of friendships at a university department. Each node corresponds to a person, and an edge indicates friendship.

*The network has been loaded as networkx graph object `G1`.*

```python
import networkx as nx

G1 = nx.read_gml('assets/friendships.gml')
```

### Question 1

Find the degree centrality, closeness centrality, and betweeness centrality of node 100.

*This function should return a tuple of floats `(degree_centrality, closeness_centrality, betweenness_centrality)`.*

<!-- grade_id: cell-bffef7d5151cb470, points: 0, grade: false, solution: true -->
```python
def answer_one():
  degCent = nx.degree_centrality(G1)
  clCent = nx.closeness_centrality(G1)
  betCent = nx.betweenness_centrality(G1)
  return (degCent[100],clCent[100],betCent[100])

answer_one()
```

Output:
```
(0.0026501766784452294, 0.2654784240150094, 7.142902633244772e-05)
```

<!-- grade_id: cell-cd6a99ae1fdd71a9, points: 1, grade: true, solution: false -->
```python
ans_one = answer_one()
assert type(ans_one) == tuple, "You must return a tuple"

```

### Use centrality measures to answer questions 2-4

### Question 2

Suppose you are employed by an online shopping website and are tasked with selecting one user in network G1 to send an online shopping voucher to. We expect that the user who receives the voucher will send it to their friends in the network.  You want the voucher to reach as many nodes as possible. The voucher can be forwarded to multiple users at the same time, but the travel distance of the voucher is limited to one step, which means if the voucher travels more than one step in this network, it is no longer valid. Apply your knowledge in network centrality to select the best candidate for the voucher.

*This function should return an integer, the chosen node.*

<!-- grade_id: cell-f6296c30200cc9d0, points: 0, grade: false, solution: true -->
```python
def answer_two():
  degCent = nx.degree_centrality(G1)
  max_node = max(degCent, key = degCent.get)
  return max_node

answer_two()
```

Output:
```
105
```

<!-- grade_id: cell-d890b05007b8cce5, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

### Question 3

Now the limit of the voucher’s travel distance has been removed. Because the network is connected, regardless of who you pick, every node in the network will eventually receive the voucher. However, we now want to ensure that the voucher reaches nodes as quickly as possible (i.e. in the fewest number of hops). How will you change your selection strategy? Write a function to tell us who is the best candidate in the network under this condition.

*This function should return an integer, the chosen node.*

<!-- grade_id: cell-e703cf046ffacac6, points: 0, grade: false, solution: true -->
```python
def answer_three():
  clCent = nx.closeness_centrality(G1)
  max_node = max(clCent, key = clCent.get)
  return max_node

answer_three()
```

Output:
```
23
```

<!-- grade_id: cell-467293083292b0ee, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

### Question 4

Assume the restriction on the voucher’s travel distance is still removed, but now a competitor has developed a strategy to remove a person from the network in order to disrupt the distribution of your company’s voucher. You competitor plans to remove people who act as bridges in the network. Identify the best possible person to be removed by your competitor?

*This function should return an integer, the chosen node.*

<!-- grade_id: cell-0aa27c6d1caf2520, points: 0, grade: false, solution: true -->
```python
def answer_four():
  betCent = nx.betweenness_centrality(G1)
  max_node = max(betCent, key = betCent.get)
  return max_node

answer_four()
```

Output:
```
333
```

<!-- grade_id: cell-338a3ca88864385b, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()

```

## Part 2

`G2` is a directed network of political blogs, where nodes correspond to a blog and edges correspond to links between blogs. Use your knowledge of PageRank and HITS to answer Questions 5-9.

```python
G2 = nx.read_gml('assets/blogs.gml')
```

### Question 5

Apply the Scaled Page Rank Algorithm to this network. Find the Page Rank of node 'realclearpolitics.com' with damping value 0.85.

*This function should return a float.*

<!-- grade_id: cell-29a0e65bd30b8324, points: 0, grade: false, solution: true -->
```python
def answer_five():
  page_rank = nx.pagerank(G2, alpha = 0.85)
  return page_rank['realclearpolitics.com']

answer_five()
```

Output:
```
0.004636694781649098
```

<!-- grade_id: cell-5ade64a5771dcbce, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Apply the Scaled Page Rank Algorithm to this network with damping value 0.85. Find the 5 nodes with highest Page Rank.

*This function should return a list of the top 5 blogs in desending order of Page Rank.*

<!-- grade_id: cell-dc31ae880c39420f, points: 0, grade: false, solution: true -->
```python
def answer_six():
  page_rank = nx.pagerank(G2, alpha = 0.85)
  sort_pr = sorted(page_rank.items(), key = lambda x: x[1])[-5:]
  return [node for node,_ in sort_pr]

answer_six()
```

Output:
```
['talkingpointsmemo.com',
 'blogsforbush.com',
 'instapundit.com',
 'atrios.blogspot.com',
 'dailykos.com']
```

<!-- grade_id: cell-fa118135cb4998f4, points: 1, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == list, "You must return a list"

```

### Question 7

Apply the HITS Algorithm to the network to find the hub and authority scores of node 'realclearpolitics.com'.

*Your result should return a tuple of floats `(hub_score, authority_score)`.*

<!-- grade_id: cell-bc1c0cf4f5f4f66a, points: 0, grade: false, solution: true -->
```python
def answer_seven():
  hist = nx.hits(G2)
  return (hist[0]['realclearpolitics.com'], hist[1]['realclearpolitics.com'])

answer_seven()
```

Output:
```
(0.0003243556140278741, 0.003918957644934261)
```

<!-- grade_id: cell-43b3de064e549ef6, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()
assert type(ans_seven) == tuple, "You must return a tuple"

```

### Question 8

Apply the HITS Algorithm to this network to find the 5 nodes with highest hub scores.

*This function should return a list of the top 5 blogs in desending order of hub scores.*

<!-- grade_id: cell-8f45f0320e6c645b, points: 0, grade: false, solution: true -->
```python
def answer_eight():
  hubs = nx.hits(G2)[0]
  sort_hubs = sorted(hubs.items(), key = lambda x: x[1], reverse=True)[:5]
  return [node for node,_ in sort_hubs]

answer_eight()
```

Output:
```
['politicalstrategy.org',
 'madkane.com/notable.html',
 'liberaloasis.com',
 'stagefour.typepad.com/commonprejudice',
 'bodyandsoul.typepad.com']
```

<!-- grade_id: cell-72499b780b38eb2c, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()
assert type(ans_eight) == list, "You must return a list"

```

### Question 9

Apply the HITS Algorithm to this network to find the 5 nodes with highest authority scores.

*This function should return a list of the top 5 blogs in desending order of authority scores.*

<!-- grade_id: cell-2c2606d7aaaae35a, points: 0, grade: false, solution: true -->
```python
def answer_nine():
  auths = nx.hits(G2)[1]
  sort_auth = sorted(auths.items(), key = lambda x: x[1], reverse=True)[:5]
  return [node for node,_ in sort_auth]

answer_nine()
```

Output:
```
['dailykos.com',
 'talkingpointsmemo.com',
 'atrios.blogspot.com',
 'washingtonmonthly.com',
 'talkleft.com']
```

<!-- grade_id: cell-bbc73cedc13c80ca, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == list, "You must return a list"

```